# Q1: Team Management System

Your task is to implement a simple team management system in Python using object-oriented programming concepts. You will create two classes: `Member` and `Team`. You will also write a main program to demonstrate how these classes can be used.

The **Member** class represents an individual team member. It should have the following:

- `__init__(name: str, role: str)`: Initializes a member with their name and role.

- `__str__()`: Returns a string representation of the member's name and role in the format "Name: {name}, Role: {role}". For example,
```python
member = Member("Alice", "Developer")
print(member)  # Name: Alice, Role: Developer
```

The **Team** class manages a collection of Member objects. It should have the following:

- `__init__(team_name: str)`: Initializes a team with its name and an empty list of members.
- `add_member(member: Member)`: Adds a Member to the team.
- `remove_member(name: str)`: Removes a Member from the team based on their name.
- `count_members_by_role(self)`: Returns a dictionary where keys are roles and values are the count of members with that role.
- `role_percentage(self)`: Returns a dictionary where keys are roles and values are the percentage of members with that role, formatted as percentages.

<!-- grade_id: cell-579f0d6ef187ca06, points: 0, grade: false, solution: true -->
```python
from collections import Counter

class Member:
    def __init__(self, name, role):
        self.name = name
        self.role = role
    def __str__(self):
        return f"Name: {self.name}, Role: {self.role}"
        
class Team:
    def __init__(self, team_name):
        self.team_name = team_name
        self.members = []

    def add_member(self, member):
        self.members.append(member)

    def remove_member(self, name):
        for i in self.members:
            if i.name == name:
                self.members.remove(i)

    def count_members_by_role(self):
        role_counts = Counter(member.role for member in self.members)
        return dict(role_counts)

    def role_percentage(self):
        dic = self.count_members_by_role()
        dic1 = {role : round(per/sum(list(dic.values())) * 100, 2) for role, per in dic.items()}
        return dic1

    def __str__(self):
        return f"Team Name: {self.team_name}, Number of Members: {len(self.members)}"
```

<!-- grade_id: cell-14a70865bc4c5464, points: 1, grade: true, solution: false -->
```python
# Test: Verify that the Member object is initialized with the correct name and role.
def test_member_constructor():
    member = Member("Alice", "Developer")
    assert member.name == "Alice"
    assert member.role == "Developer"
    print("ok")

test_member_constructor()
```

Output:
```
ok
```

<!-- grade_id: cell-03181b05a8ae7a22, points: 1, grade: true, solution: false -->
```python
# Test: Verify that the string representation of the Member object is correctly formatted.
def test_member_str():
    member = Member("Bob", "Designer")
    assert str(member) == "Name: Bob, Role: Designer"
    print("ok")

test_member_str()
```

Output:
```
ok
```

<!-- grade_id: cell-70c3d77d3f129c51, points: 1, grade: true, solution: false -->
```python
# Test: Verify that a member is added to the team.
def test_add_member():
    team = Team("Project A")
    member = Member("Alice", "Developer")
    team.add_member(member)
    assert len(team.members) == 1
    assert team.members[0].name == "Alice"
    print("ok")

test_add_member()
```

Output:
```
ok
```

<!-- grade_id: cell-6718d481ca6b2e78, points: 1, grade: true, solution: false -->
```python
# Test: Verify that a member is removed from the team.
def test_remove_member():
    team = Team("Project A")
    member1 = Member("Alice", "Developer")
    member2 = Member("Bob", "Designer")
    team.add_member(member1)
    team.add_member(member2)
    team.remove_member("Alice")
    assert len(team.members) == 1
    assert team.members[0].name == "Bob"
    print("ok")

test_remove_member()
```

Output:
```
ok
```

<!-- grade_id: cell-07c7a345aeacbe30, points: 1, grade: true, solution: false -->
```python
# Test: Verify that the number of members is counted correctly by role.
def test_count_members_by_role():
    team = Team("Project A")
    team.add_member(Member("Alice", "Developer"))
    team.add_member(Member("Bob", "Designer"))
    team.add_member(Member("Charlie", "Developer"))
    role_counts = team.count_members_by_role()
    assert role_counts == {"Developer": 2, "Designer": 1}
    print("ok")

test_count_members_by_role()
```

Output:
```
ok
```

<!-- grade_id: cell-1665ebeb7c3acb5f, points: 1, grade: true, solution: false -->
```python
# Test: Verify that the percentage of members for each role is calculated correctly.
def test_role_percentage():
    team = Team("Project A")
    team.add_member(Member("Alice", "Developer"))
    team.add_member(Member("Bob", "Designer"))
    team.add_member(Member("Charlie", "Developer"))
    percentages = team.role_percentage()
    assert abs(percentages["Developer"] - 66.67) < 0.01
    assert abs(percentages["Designer"] - 33.33) < 0.01
    print("ok")

test_role_percentage()
```

Output:
```
ok
```

```python
# Example usage
# Create some team members
member1 = Member("Alice", "Developer")
member2 = Member("Bob", "Designer")
member3 = Member("Charlie", "Developer")
member4 = Member("Diana", "Designer")
member5 = Member("Eve", "Project Manager")

# Create a team and add members
team = Team("Project A")
team.add_member(member1)
team.add_member(member2)
team.add_member(member3)
team.add_member(member4)
team.add_member(member5)

# Print team details
print(team)

# Calculate and print role counts
print("\nRole Counts:")
role_counts = team.count_members_by_role()
for role, count in role_counts.items():
    print(f"{role}: {count}")

# Calculate and print role percentages
print("\nRole Percentages:")
percentages = team.role_percentage()
for role, percentage in percentages.items():
    print(f"{role}: {percentage:.2f}%")
```

Output:
```
Team Name: Project A, Number of Members: 5

Role Counts:
Developer: 2
Designer: 2
Project Manager: 1

Role Percentages:
Developer: 40.00%
Designer: 40.00%
Project Manager: 20.00%
```

```python

```